# Exercise: Root Finding 1
**강좌**: *Numerical Analysis*

**Due**: 2024/10/1

## Problem #1
Write a function to find a root using the Bisection method with the following specifications:
- Use a for loop with a specified maximum number of iterations.
- Input parameters
   * start and end of the braket (*a* and *b*)
   * function *f*
   * tolerance *tol*
   * maximum iterations *iter_max*
   * (Optional) Boolean verbose keyword to print intermediate results. The default is False.
- Additionally, include a proper docstring to describe the function's behavior, parameters, and return values.

<!-- grade_id: cell-fe271d986d99dd1f, points: 0, grade: false, solution: true -->
```python
# Write your answer
def bisect(a, b, f, tol=1e-6, iter_max=50):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-962d2cf9716a6461, points: 5, grade: true, solution: false -->
```python
import numpy as np
f = lambda x: x**2 + 10*np.sin(x)

assert abs(bisect(-10, -1,  f) - (-2.47948)) < 5e-5
assert abs(bisect(-1, 1,  f) - (0.0)) < 1e-6
assert abs(bisect(-1, 0,  f) - (0.0)) < 1e-6
```

## Problem #2
Write a function to find a root using the False position method with the following specifications:
- Use a for loop with a specified maximum number of iterations.
- Input parameters
   * start and end of the braket (*a* and *b*)
   * function *f*
   * tolerance *tol*
   * maximum iterations *iter_max*
   * (Optional) Boolean verbose keyword to print intermediate results. The default is False.
- Additionally, include a proper docstring to describe the function's behavior, parameters, and return values.

<!-- grade_id: cell-604cece7f9e1450f, points: 0, grade: false, solution: true -->
```python
# Write your answer
def falseposition(a, b, f, tol=1e-6, iter_max=50):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f84270eb85482c08, points: 5, grade: true, solution: false -->
```python
import numpy as np
f = lambda x: x**2 + 10*np.sin(x)

assert abs(falseposition(-10, -1,  f) - (-2.47948)) < 5e-5
assert abs(falseposition(-1, 1,  f) - (0.0)) < 1e-6
assert abs(falseposition(-1, 0,  f) - (0.0)) < 1e-6
```

## Problem #3
Determine the real root of $f(x)=(0.8-0.3x)/x$

(a) Graphically.

(b) Bisection method

(c) False position method

For (b) and (c) use initial guesses of $a=1$ and $b=3$, and iterate until the approximate error falls below $0.1%$ or the number of iterations exceeds 50. 

<!-- grade_id: cell-f6c57147fce4c063, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-57a9b7c6abe35785, points: 0, grade: false, solution: true -->
```python
# Write your answers (b)
# Save the result in ans_b

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5b6fd022369d37ae, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-ccc08edda3d585d2, points: 0, grade: false, solution: true -->
```python
# Write your answers (c)
# Save the result in ans_c

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-01ff38e473ed2c83, points: 2, grade: true, solution: false -->
```python

```

## Problem #4
Water is flowing in a trapezoidal channel at a rate of $Q=20m^3/s$. The critical depth y for such a channel must satisfy the equation
$$
0 = 1 - \frac{Q^2}{gA_c^3} B
$$
where $g=9.81 m/s2$, $A_c =$ the cross-sectional area ($m^2$), and $B=$ the width of the channel at the surface (*m*). For this case, the width and the cross-sectional area can be related to depth *y* by
$$
B = 3 + y, A_c  = 3y + \frac{y^2}{2}
$$
Solve for the critical depth using 
(a) the graphical method, (b) bisection, and (c) false position. 
For (b) and (c) use initial guesses of $a=0.5$ and $b=2.5$, and iterate until the approximate error falls below $0.1 \%$. 

<!-- grade_id: cell-9b7cadf3621cc1ce, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7e4cca84c8892f74, points: 0, grade: false, solution: true -->
```python
# Write your answers (b)
# Save the result in ans_b

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-19245e5c403472e1, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-0a8e1726660934e0, points: 0, grade: false, solution: true -->
```python
# Write your answers (c)
# Save the result in ans_c

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-66cb9bdb83a8bfd4, points: 2, grade: true, solution: false -->
```python

```