# Exercise: Root Finding 2
**강좌**: *Numerical Analysis*

**Due**: 2024/10/8

## Problem #1
Write a function to find a root using the Newton method with the following specifications:
- Use a for loop with a specified maximum number of iterations.
- Input parameters
   * function *f*
   * derivative *fp*
   * intial guess: x0
   * tolerance *tol*
   * maximum iterations *iter_max*
   * (Optional) Boolean verbose keyword to print intermediate results. The default is False.
- Additionally, include a proper docstring to describe the function's behavior, parameters, and return values.

<!-- grade_id: cell-fe271d986d99dd1f, points: 0, grade: false, solution: true -->
```python
# Write your answer
def newton(f, fp, x0, tol=1e-6, iter_max=50):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-962d2cf9716a6461, points: 5, grade: true, solution: false -->
```python
import numpy as np
f = lambda x : x**2 + 10*np.sin(x)
fp = lambda x: 2*x + 10*np.cos(x)

assert abs(newton(f, fp, x0=-4) - (-2.47948)) < 5e-5
assert abs(newton(f, fp, x0=-1) - (0.0)) < 1e-6
assert abs(newton(f, fp, x0=0)- (0.0)) < 1e-6
```

## Problem #2
Write a function to find a root using the Secant method with the following specifications:
- Use a for loop with a specified maximum number of iterations.
- Input parameters
   * function *f*
   * intial guesses: $x_0$, $x_{-1}$
   * tolerance *tol*
   * maximum iterations *iter_max*
   * (Optional) Boolean verbose keyword to print intermediate results. The default is False.
- Additionally, include a proper docstring to describe the function's behavior, parameters, and return values.

<!-- grade_id: cell-604cece7f9e1450f, points: 0, grade: false, solution: true -->
```python
# Write your answer
def secant(f, x0, x1, tol=1e-6, iter_max=50):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f84270eb85482c08, points: 5, grade: true, solution: false -->
```python
f = lambda x : x**2 + 10*np.sin(x)

assert abs(secant(f, x0=-4, x1=-5) - (-2.47948)) < 5e-5
assert abs(secant(f, x0=-1, x1=-1.1) - (0.0)) < 1e-6
assert abs(secant(f, x0=0, x1=-1)- (0.0)) < 1e-6
```

# Problem #3
Determine the highest real root of

$$
f(x) = 2x^3 - 11.7x^2 + 17.7x - 5
$$

(a) Graphically.

(b) Newton-Raphson method with $x_0 = 3$.

(c) Secant method with $x_{-1}=3, x_0=4$.

<!-- grade_id: cell-f6c57147fce4c063, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-57a9b7c6abe35785, points: 0, grade: false, solution: true -->
```python
# Write your answers (b)
# Save the result in ans_b

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5b6fd022369d37ae, points: 2, grade: true, solution: false -->
```python
# Do not remove!!!
```

<!-- grade_id: cell-ccc08edda3d585d2, points: 0, grade: false, solution: true -->
```python
# Write your answers (c)
# Save the result in ans_c

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-01ff38e473ed2c83, points: 2, grade: true, solution: false -->
```python
# Do not remove!!!
```

## Problem #4
Aerospace engineers sometimes compute the trajectories of projectiles like rockets. A related problem deals with the trajectory of a
thrown ball. The trajectory of a ball is defi ned by the (x, y) coordinates. The trajectory can be modeled as
$$
y = \tan(\theta_0)x - \frac{g}{2v_0^2\cos^2\theta_0} x^2+y_0
$$
Find the appropriate initial angle $\theta_0$, if the initial velocity $v_0=20m/s$ and the disctance to the catcher $x$ is $40m$. Note that the ball leaves the thrower’s hand at an elevation of $y_0=1.8 m$ and and the catcher receives it at $1m$. Express the fi nal result in degrees. Use a value of $9.81m/s^2$ for $gG.

(a) Graphically and determine the initial guess

(b) Secant method

<!-- grade_id: cell-9b7cadf3621cc1ce, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7e4cca84c8892f74, points: 0, grade: false, solution: true -->
```python
# Write your answers (b)
# Save the result in ans_b

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-19245e5c403472e1, points: 3, grade: true, solution: false -->
```python
# Do not remove!!!
```