Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Exercise: Regression, Interpolation
**강좌**: *Numerical Analysis*

**Due**: 2024/11/20

```python
%matplotlib inline
from matplotlib import pyplot as plt

import numpy as np

plt.style.use('ggplot')
plt.rcParams['figure.dpi'] = 150
```

## Problem #1
Write a function to perform nn-th order regression analysis for a given dataset.

- The function should return the regression coefficients.
- The function should print the correlation coefficient.

<!-- grade_id: cell-fe271d986d99dd1f, points: 0, grade: false, solution: true -->
```python
# Write your answer
def polyreg(x, y, n):
    """
    Regression for polynomial
    
    Parameters
    ----------
    x : array
        x
    y : array
        y
    n : integer
        order
    
    Returns
    --------
    z : array
        coeffcient of polynomials
    r : float
        correlation coefficient
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-104bcd844790831f, points: 5, grade: true, solution: false -->
```python
## Do not remove
# Same problem of lecture note
y = np.array([2.1, 7.7, 13.6, 27.2, 40.9, 61.1])
x = np.array(range(6))
z, r = polyreg(x, y, 2)

# Check your result
assert np.linalg.norm(z - np.array([2.47857143, 2.35928571, 1.86071429])) < 1e-6
```

## Problem #2
Write a function to perform interpolation using the Lagrange polynomial.:

- Using the function values of $log(x)$ at $x=1,3,4,5,6$, calculate the interpolated value at $x=2$.

- Solve the same problem using Newton's divided difference method and compare the results.

<!-- grade_id: cell-e6fa2705e67c6109, points: 0, grade: false, solution: true -->
```python
def lagint(xs, fs, x):
    """
    Lagrange Interpolation
    
    Parameters
    ----------
    xs : array
        Data of x
    ys : array
        Data of y
    x: float
        Point to interpolate
    """
    # YOUR CODE HERE
    raise NotImplementedError()


# Interpolate log(2) using divided-difference interpolating polynomial using x=1,3,4,5,6 (save as fdd)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e1ccf1c13c7ac806, points: 5, grade: true, solution: false -->
```python
## Do not remove
xs = [1, 3, 4, 5, 6]
fs = np.log(xs)

# Interpolate using Lagrange Polynomial
flg = lagint(xs, fs, 2)

# Check your result
assert abs(flg - fdd) < 1e-6
```

## Problem #3
Hooke’s law, which holds when a spring is not stretched too far, signifies that the extension of the spring and the applied force are linearly related. The proportionality is parameterized by the spring constant $k$. A value for this parameter can be established experimentally by placing known weights onto the spring and measuring the resulting compression. Such data were contained in below table and plotted in below figure. Notice that above a weight $40\times10^4 N$, the linear relationship between the force and displacement breaks down. This sort of behavior is typical of what is termed a “hardening spring.” 

(a) Employ linear regression to determine a value of k for the linear portion of this system. Fit a quadratic polynomial to the nonlinear portion. (yellow region)

| Displacement, m | 0.10 | 0.17 | 0.27 | 0.35 | 0.39 | 0.42 | 0.43 | 0.44 |
|-----------------|------|------|------|------|------|------|------|------|
| Force, $10^4 N$ | 10 | 20 | 30 | 40 | 50 | 60 | 70 | 80 |

```python
x = np.array([0.10, 0.17, 0.27, 0.35, 0.39, 0.42, 0.43, 0.44])
f = np.array([10, 20, 30, 40, 50, 60, 70, 80])

plt.plot(x, f, linestyle='none', marker='o')
plt.axvspan(0.35, 0.5, alpha=0.2, color='yellow')
plt.xlim(0.05, 0.5)
plt.xlabel('Displacement, m')
plt.ylabel('Force, $10^4$ N')
```

<!-- grade_id: cell-06d903be04fc19c3, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
# Quadratic interpolation function name should be saved as 'fhard'.
raise NotImplementedError()
```

<!-- grade_id: cell-eddbe11c8ecd0bba, points: 4, grade: true, solution: false -->
```python
## Do not remove
x = np.array([0.10, 0.17, 0.27, 0.35, 0.39, 0.42, 0.43, 0.44])
y = fhard(x)

# Check your result
assert abs(f[1] - y[1]) < 182
assert abs(f[3] - y[3]) < 0.6
assert abs(f[6] - y[6]) < 0.3
```

## Problem #4
Runge’s function is written as:


$$
f(x) = \frac{1}{1 + 25x^2}
$$

(a) Develop a plot of this function for the interval from x = -1 to 1.

(b) Generate and plot the fourth-order Lagrange interpolating polynomial using equispaced function values corresponding to x = -1, -0.5, 0, 0.5, and 1.


<!-- grade_id: cell-bd8a7af00b1157ca, points: 1, grade: true, solution: true -->
```python
# Write your answer (a)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0923faedc04f947c, points: 1, grade: true, solution: true -->
```python
# Write your answer (b)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-14b73e1e903a2427, points: 3, grade: true, solution: false -->
```python
# Do not remove!!
def Runge(x):
    return 1/(1 + 25*x**2)

xis = np.array([-1, -0.5, 0, 0.5, 1])
fis = Runge(xis)
lag = lagint(xis, fis, xis)

```