Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Exercise: Gauss Eliniation, LU Decomposition
**강좌**: *Numerical Analysis*

**Due**: 2024/10/21

## Problem #1
Write a function to solve $Ax=b$ using Gauss Elimination with Partial Pivoting. The function should follow these specifications:

- Input parameters
   * Matrix *A*
   * vector *b*
   
- Return parameters
   * solution vector *x*
   * Pivot indices *l* (indicating row swaps during pivoting)
   
- Additionally, include a proper docstring to describe the function's behavior, parameters, and return values.

<!-- grade_id: cell-fe271d986d99dd1f, points: 0, grade: false, solution: true -->
```python
import numpy as np


def pv_gauss_eliminate(A, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-104bcd844790831f, points: 5, grade: true, solution: false -->
```python
## Do not remove
A = np.array([
    [3, -13, 9, 3],
    [-6, 4, 1, -18],
    [6, -2, 2, 4],
    [12, -8, 6, 10]
])

b = np.array([-19, -34, 16, 26])
x, l = pv_gauss_eliminate(A, b)

# Check your result
assert np.linalg.norm(x - np.array([3, 1, -2, 1])) < 1e-6
assert np.linalg.norm(l - np.array([3, 0, 1, 2])) < 1e-6
```

## Problem #2
Write two functions to perform LU decomposition ($A=LU$) and substitutions to solve $Ax=b$ using forward and backward substitutions with the following specifications:

- LU decomposition function
    * Convert matrix A into its LU decomposed form.
    
- LU substitution function
    * Input parameters
       * Matrix *LU* (from the LU decomposition)
       * vector *b* (right-hand side of the equation)
       
    - Return parameters
       * solution vector *x* (final solution to $Ax=b$)
       * Intermediate result of forward substitution $Lc=b$
       
- Additionally, include a proper docstring to describe the function's behavior, parameters, and return values.

<!-- grade_id: cell-e6fa2705e67c6109, points: 0, grade: false, solution: true -->
```python
def lu_decompose(A):
    # YOUR CODE HERE
    raise NotImplementedError()

def lu_subs(A, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e1ccf1c13c7ac806, points: 5, grade: true, solution: false -->
```python
## Do not remove
A = np.array([[2, 1, 1], [4, -6, 0], [-2, 7, 2]])
LU = A.copy()

# LU decomposition
lu_decompose(LU)

# Solve Ax=b
b = np.array([5, -2, 9])
x, c = lu_subs(LU, b)

# Check your result
assert np.linalg.norm(LU - np.array([[2, 1, 1], [2, -8, -2], [-1, -1, 1]])) < 1e-6
assert np.linalg.norm(x - np.array([1, 1, 2])) < 1e-6
assert np.linalg.norm(c - np.array([5, -12, 2])) < 1e-6
```

## Problem #3
Write a function to compute the inverse of matrix AA using LU decomposition and substitution.

<!-- grade_id: cell-06d903be04fc19c3, points: 0, grade: false, solution: true -->
```python
def inverse(A):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-26344620df2a1a57, points: 0, grade: true, solution: false -->
```python
## Do not remove
A = np.array([[2,1,1],[4,-6,0],[-2,7,2]])
A_inv = inverse(A)

# Check your result
assert np.linalg.norm(A @ A_inv - np.eye(len(A))) < 1e-6
```

## Problem #4
Write a function to solve a tri-diagonal system using the Thomas algorithm to solve the following problem:

$$
\begin{bmatrix}
2.01475 & -0.020875 & & \\
-0.020875 & 2.01475 & -0.020875 & \\
& -0.020875 & 2.01475 & -0.020875 \\
& & -0.020875 & 2.01475
\end{bmatrix}
T 
= 
\begin{bmatrix}
4.175 \\ 0 \\ 0 \\ 2.0875
\end{bmatrix}
$$

<!-- grade_id: cell-0923faedc04f947c, points: 0, grade: false, solution: true -->
```python
def thomas(a, b, c, d):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7db0ce8d66c18a84, points: 0, grade: false, solution: true -->
```python
# Solve the given linear system, store the solution in the vector x.
# print(x)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3913d02b393042c3, points: 5, grade: true, solution: false -->
```python
# Do not remove!!!
```