Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

```python
%matplotlib inline
from matplotlib import pyplot as plt

import numpy as np

plt.style.use('ggplot')
plt.rcParams['figure.dpi'] = 150
```

# Exercise: Initial Value Problem and Boundary Value Problem
**강좌**: *Numerical Analysis*

**Due**: 2024/12/19

## Problem 1  

Write code to solve the initial value problem (for a scalar variable) using the following numerical methods:  

1. Explicit Euler Method  
2. Second-Order Runge-Kutta Method (Richardson)  
3. Fourth-Order Runge-Kutta Method  

Additionally, use the implemented methods to solve the **paratrooper problem** presented in the lecture notes.  

### Requirements:  
- Implement each method in Python as a separate function.  
- Test your implementations on the paratrooper problem with the given initial conditions and parameters.  
- Compare the numerical solutions by varying the time step ($h$) as follows:  
  - $h = 0.5, 1.0, 2.0, 5.0, 10.0, 15.0$
- Discuss how the choice of time step affects the accuracy and stability of the solutions.  

<!-- grade_id: cell-fe271d986d99dd1f, points: 0, grade: false, solution: true -->
```python
# Write Explicit Euler
def explicit_euler(f, tspan, y0, h):
    """
    Explicit Euler Method
    
    Parameter
    ---------
    f : function
        Derivative
    tspan : tuple
        Initial and final time ex) (ti, te)
    y0 : float
        Initial solution
    h: float
        Time step size
        
    Return
    ------
    t : array
        Time series
    y : array
        solutions
    """
    # YOUR CODE HERE

    return t, y
```

<!-- grade_id: cell-ffb7da9baaf80d92, points: 0, grade: false, solution: true -->
```python
# Write 2nd-order Runge Kutta
def runge_kutta2(f, tspan, y0, h):
    """
    Explicit Euler Method
    
    Parameter
    ---------
    f : function
        Derivative
    tspan : tuple
        Initial and final time ex) (ti, te)
    y0 : float
        Initial solution
    h: float
        Time step size
        
    Return
    ------
    t : array
        Time series
    y : array
        solutions
    """
    # YOUR CODE HERE

    return t, y
```

<!-- grade_id: cell-3a1cc593f678b4ce, points: 0, grade: false, solution: true -->
```python
# Write 4th-order Runge Kutta
def runge_kutta4(f, tspan, y0, h):
    """
    Explicit Euler Method
    
    Parameter
    ---------
    f : function
        Derivative
    tspan : tuple
        Initial and final time ex) (ti, te)
    y0 : float
        Initial solution
    h: float
        Time step size
        
    Return
    ------
    t : array
        Time series
    y : array
        solutions
    """
    # YOUR CODE HERE

    return t, y
```

```python
import numpy as np

# Paratrooper problem
m, c, g = 68.1, 12.5, 9.81

# function
def f(t, v):
    return g - c/m*v

# Time step
h = 2.0
eet, eey = explicit_euler(f, (0, 50), 0, h)
r2t, r2y = runge_kutta2(f, (0, 50), 0, h)
r4t, r4y = runge_kutta4(f, (0, 50), 0, h)

# Verification
# Plot exact solution and numerical solution


```

<!-- grade_id: cell-2da89e50fbcfa95b, points: 5, grade: true, solution: false -->
```python
# Check your solution
assert r4y[25] - 53.439347 < 1e-6
```

## Problem 2  

Write code to solve the initial value problem (for a vector variable) using the following numerical methods:  

1. Explicit Euler Method  
2. Second-Order Runge-Kutta Method (Richardson)  
3. Fourth-Order Runge-Kutta Method  

Use the implemented methods to solve the **spring-mass problem** described in the lecture notes.  

### Requirements:  
- Implement each method in Python as a separate function to handle vector-valued problems.  
- Apply the methods to solve the spring-mass problem with the given initial conditions and parameters. 
- Compare the numerical solutions by varying the time step ($h$) as follows:  
  - $h = 0.01, 0.05, 0.1, 0.2$
- Compare the results for the different methods, analyzing accuracy and stability.  

<!-- grade_id: cell-bce534231aa2aea3, points: 0, grade: false, solution: true -->
```python
# Write Explicit Euler (for vector y)
def explicit_euler(f, tspan, y0, h):
    """
    Explicit Euler Method
    
    Parameter
    ---------
    f : function
        Derivative
    tspan : tuple
        Initial and final time ex) (ti, te)
    y0 : float
        Initial solution
    h: float
        Time step size
        
    Return
    ------
    t : array
        Time series
    y : array
        solutions
    """
    # YOUR CODE HERE

    return t, y
```

<!-- grade_id: cell-788bf67c14c058a7, points: 0, grade: false, solution: true -->
```python
# Write 2nd-order Runge Kutta (for vector y)
def runge_kutta2(f, tspan, y0, h):
    """
    Explicit Euler Method
    
    Parameter
    ---------
    f : function
        Derivative
    tspan : tuple
        Initial and final time ex) (ti, te)
    y0 : float
        Initial solution
    h: float
        Time step size
        
    Return
    ------
    t : array
        Time series
    y : array
        solutions
    """
    # YOUR CODE HERE

    return t, y
```

<!-- grade_id: cell-e563507ed654ca8c, points: 0, grade: false, solution: true -->
```python
# Write 4th-order Runge Kutta (for vector y)
def runge_kutta4(f, tspan, y0, h):
    """
    Explicit Euler Method
    
    Parameter
    ---------
    f : function
        Derivative
    tspan : tuple
        Initial and final time ex) (ti, te)
    y0 : float
        Initial solution
    h: float
        Time step size
        
    Return
    ------
    t : array
        Time series
    y : array
        solutions
    """
    # YOUR CODE HERE

    return t, y
```

```python
# Paratrooper problem
m, k = 1, 16

# Derivative
def f(t, y):
    return np.array([y[1], -k/m*y[0]])

# Initial condition
y0 = np.array([1, 0])

# Time step
h = 0.05
eet, eey = explicit_euler(f, (0, 5), y0, h)
r2t, r2y = runge_kutta2(f, (0, 5), y0, h)
r4t, r4y = runge_kutta4(f, (0, 5), y0, h)

# Verification
# Plot exact solution and numerical solution

```

<!-- grade_id: cell-37c27501552fe516, points: 5, grade: true, solution: false -->
```python
# Check your solution
assert r4y[0, 100] - 0.408303 < 1e-6
```

## Problem 3  

A forced damped spring-mass system is described by the following nonlinear ordinary differential equation of motion:  

$$
m \frac{d^2 x}{dt^2} + a \left| \frac{dx}{dt} \right| \frac{dx}{dt} + kx = F_o \sin (\omega t)
$$  

where:  
- $x(t)$: Displacement from equilibrium.  
- $t$: Time.  
- $m = 2 \, \text{kg} $: Mass.  
- $ a = 5 \, \text{N}/(\text{m/s})^2 $: Damping coefficient (nonlinear air damping).  
- $ k = 6 \, \text{N/m} $: Spring stiffness.  
- $F_o = 2.5 \, \text{N} $: Amplitude of the forcing function.  
- $ \omega = 0.5 \, \text{rad/sec} $: Angular frequency of the forcing function.  

### Initial Conditions  
- Initial velocity: $ \frac{dx}{dt} = 0 \, \text{m/s} $.  
- Initial displacement: $ x = 1 \, \text{m} $.  

### Task  
Solve the equation numerically using the **4th-order Runge-Kutta method** over the time period $0 \leq t \leq 15 \, \text{s} $.  


<!-- grade_id: cell-5683929632f1c3db, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE


# Verification
# Plot exact solution and numerical solution
```

<!-- grade_id: cell-bc3d28c6459e143d, points: 5, grade: true, solution: false -->
```python
# Check your solution
assert r4y[0, 100] - 0.711351 < 1e-6
```

## Problem 4
Repeat the boundary value problem (BVP) example provided in the lecture notes using the **shooting method**. Implement the **4th-order Runge-Kutta method** instead of using the `solve_ivp` function.  

<!-- grade_id: cell-d134483c7b632131, points: 0, grade: false, solution: true -->
```python
from scipy.optimize import root_scalar


# Constants
Ta, T1, T2 = 20, 40, 200
h = 0.01

# derivative function
def dydx(t, y):
    # y[0], y[1] : T and z
    return np.array([y[1], h*(y[0] - Ta)])

# YOUR CODE HERE

```

<!-- grade_id: cell-1423b2b37a373b61, points: 5, grade: true, solution: false -->
```python
# Check your solution
assert y[1, 1000] - 21.932798 < 1e-6
```

## Problem 5
The boundary layer velocity profile can be obtained by solving the Blasius equation:  

$$
2 f''' +f f'' =0
$$

### Boundary Conditions 
- $f(0) = 0, f'(0)=0$
- $f'(\infty)=1$

Using the shooting method, solve this problem by assuming the initial condition of $f''(0)$.

### Tasks  
1. **Problem Setup**:  
   - Solve the equation using the **shooting method** on $x\in[0, 10]$.  
   - Assume an initial guess for $f''(0)$ and iteratively refine it to satisfy the boundary condition $f'(10) = 1$.   

2. **Analysis**:  
   - Compute and plot the velocity profile $f'(x)$.  
   - Determine the boundary layer thickness, defined as the location where $f'(x)$ reaches 99% of the free-stream velocity (i.e., $f'(x) = 0.99$).  

<!-- grade_id: cell-2e4412c90dc501b1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

# Plot numerical solution save the result as f
# ex) x, f = runge_kutta4
```

<!-- grade_id: cell-34ccc7fb300c80e8, points: 5, grade: true, solution: false -->
```python
# Check your solution
assert np.linalg.norm(f[:, 100] - [0.16557173, 0.32978003, 0.32300712]) < 1e-6
```