Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Exercise: Integration and Differentiation
**강좌**: *Numerical Analysis*

**Due**: 2024/12/5

## Problem #1

Write Python code to implement the following numerical integration methods:  
- Develop a function to approximate the integral of a given function using the composite trapezoidal rule.  
- Develop a function to approximate the integral of a given function using Simpson's 1/3 rule.  

- Input parameters 
  * Function *f* to be integrated.  
  * Integration bounds *a* (lower limit) and *b* (upper limit).  
  * Number of subintervals *n* (must be even for Simpson's rule).
    
- Additionally, include a proper docstring to describe the function's behavior, parameters, and return values.

<!-- grade_id: cell-fe271d986d99dd1f, points: 0, grade: false, solution: true -->
```python
# Write trapezoid rule
def trapezoid(f, a, b, n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-49213c5fdfea3f31, points: 0, grade: false, solution: true -->
```python
# Write Simpson 1/3 rule
def simpson(f, a, b, n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-962d2cf9716a6461, points: 5, grade: true, solution: false -->
```python
import numpy as np

# test problem
f = np.cos
a, b = 0, np.pi/2

# exact integration on [a, b]
Ie = 1.0

# Verification
assert(abs(trapezoid(f, a, b, 10) - Ie) < 1e-2)
assert(abs(simpson(f, a, b, 10) - Ie) < 1e-5)
```

## Problem #2
Derive the finite difference formula to approximate the first derivative $f'(x)$ using the function values at points $f_{i+2}$, $f_{i+1}$, and $f_i$.  

1. **Derivation:**  
   - Use Taylor series expansions to derive the finite difference formula for $f'(x)$ with the best possible accuracy.
   - Clearly show the step-by-step process.  

2. **Implementation:**  
   - Write a Python function to compute the first derivative at $x_i$ using the derived formula.  

3. **Validation:**  
   - Test your formula and implementation by approximating the derivative of $f(x) = \sin(x)$ at $x = 0$ using a step size $h = 0.1, 0.05, 0.025$.  
   - Compare the result with the exact derivative, $f'(x) = \cos(x)$.  

<!-- grade_id: cell-ba1acc02f84c25e3, points: 0, grade: false, solution: true -->
```python
# Write your answer
def df(f, x, h):
    """
    f : function
    x : float
    h : float
        step size
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f5db8024c7703f97, points: 5, grade: true, solution: false -->
```python
# Check your result
hs = [0.1, 0.05, 0.25]
f = np.sin

assert abs(df(f, 0, hs[0]) - 1.003321) < 1e-6
assert abs(df(f, 0, hs[1]) - 1.000832) < 1e-6
assert abs(df(f, 0, hs[2]) - 1.020380) < 1e-6
```

## Problem #3
Recall that the velocity of the freefalling parachutist with linear drag can be computed analytically as

$v(t) = \frac{gm}{c}(1 - e^{-(c/m)t})$

where $v(t)$ = velocity (m/s), $t$ = time (s), $g$ = 9.81 m/s2, $m$ = mass (kg), $c$ = linear drag coefficient (kg/s). 

Use Romberg integration to compute how far the jumper travels during the first 8 seconds of free fall given $m$ = 80 kg and $c$ = 10 kg/s. 

Compute the answer to $\epsilon_{t}$ =1%.

<!-- grade_id: cell-59aa516fedc2e0e4, points: 0, grade: false, solution: true -->
```python
# Write your answer
# Save the result in ans_a
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-65087bdeedebcecc, points: 5, grade: true, solution: false -->
```python
# Do not remove!!
```

## Problem #4

The objective of this problem is to compare first-order accurate forward, backward, and second-order centered finite-difference approximations of the second derivative of a function to the actual value of the derivative. 

This will be done for

$f(x)=e^{-2x}-x$

- (a) Use calculus to determine the correct value of the second derivative at $x$ = 2.
- (b) To evaluate the second-derivative centered finite-difference approximations, start with  $\Delta x = 0.1$. Thus, for the first evaluation, the x values for the centered difference approximation will be $x$ = 1.9 and 2.1. Then, decrease in increments of 0.01 down to a minimum value of $\Delta x$ = 0.01.

- (c) Repeat (b) the second-derivative forward and backward differences. (Note that these can be done at the same time that the centered difference is computed in the loop.)
- (d) Plot the results of (b) and (c) versus x. Include the exact result on the plot for comparison.

<!-- grade_id: cell-a6e108c7a1bd00c6, points: 0, grade: false, solution: true -->
```python
# Write your answer (a)
# Save the result in ans_a
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-db74c1360ea28a7a, points: 2, grade: true, solution: false -->
```python
# Do not remove!!
```

<!-- grade_id: cell-ef8c319d03de4575, points: 3, grade: true, solution: true -->
```python
# Write your answer (d)
# YOUR CODE HERE
raise NotImplementedError()
```